# Hausaufgabe 1
## Lärmbelastung am Heimarbeitsplatz
<!-- Lizensiert unter (CC BY 4.0) Gert Herold, 2022 -->

Nach Beschwerden zahlreicher im Homeoffice arbeitender Menschen über die durch eine Dauerbaustelle auftretende Lärmbelastung wurde über mehrere Monate an Werktagen der Mittelungspegel $L_\text{A,eq}$ an einem repräsentativen Immissionsort bestimmt und mit einer (übertriebenen) Genauigkeit von 4 Nachkommastellen abgespeichert.

```python
# Werte in dB
LAeq = {
 'Januar': [60.5380, 64.6714, 53.8056, 57.6829, 60.4964, 71.7349, 67.9685, 59.9631, 64.8179, 62.0253, 55.7925,
            56.8596, 55.7310, 60.5562, 61.7910, 52.0534, 66.3174, 61.1856, 60.1812, 59.3590, 68.9449,
            60.0561],
 'Februar': [60.0914, 64.5895, 60.9965, 64.9108, 61.4350, 61.3313, 67.0028, 60.6917, 56.4955, 54.7038,
             61.2832, 56.1689, 64.7428, 62.3909, 59.4462, 57.7305, 60.6353, 63.9597, 65.2880, 55.0291],
 'März': [70.5077, 57.7741, 55.0890, 68.7610, 67.0816, 65.6817, 57.6222, 66.3466, 60.4513, 54.0338, 57.6425,
          54.6896, 65.5107, 66.8681, 55.8035, 60.2179, 62.9730, 55.1686, 65.6853, 55.9769, 54.0462],
 'April': [56.7686, 59.1396, 56.7025, 67.0300, 61.0094, 53.8614, 62.7165, 53.8287, 68.6539, 54.9897, 58.7724,
           61.9576, 57.1081, 56.0733, 62.2637, 53.1944, 66.2013, 54.5371, 54.7231, 60.2759],
 'Mai': [49.9417, 60.9106, 57.3518, 62.6660, 60.4577, 65.4958, 56.5001, 64.9796, 60.9883, 58.4091, 62.0693,
         57.1189, 58.0387, 58.3027, 61.8890, 60.4136, 50.3403, 68.1772, 69.6345],
 'Juni': [57.7947, 65.1707, 63.7792, 58.8610, 59.9703, 66.3456, 52.9550, 65.0802, 67.2341, 57.7997, 68.3582,
          64.9668, 66.4794, 60.5284, 55.8235, 59.1178, 49.6983, 56.8479, 64.3189, 56.7354]
}
```

### 1) Ausgabe für den Messbericht

**Geben Sie für jeden Monat den jeweils höchsten Pegel mit einer Nachkommastelle als Fließtext an.**
 
Beispiel:

    Im August wurde ein maximaler Mittelungspegel von 71.2 dB gemessen.
    
*(Für einen Hinweis zur Textformatierung unter Python den ausgeblendeten Abschnitt öffnen)*

```python
# Hinweis: Zeichketten in Python mit formatierenden Platzhaltern -- verschiedene Wege zum gleichen Ziel
s = 'Sekunde'
m = 'Minute'
h = 'Stunde'
s_per_m = 60
h_per_m = 1/60

print('%04d: Eine %s hat %d %sn. Das sind %8.4f %sn. Eine %s entspricht demnach %.6f %sn.' % (1, m, s_per_m, s, h_per_m, h, s, h_per_m/s_per_m, h))
print('{:04d}: Eine {} hat {} {}n. Das sind {:8.4f} {}n. Eine {} entspricht demnach {:.6f} {}n.'.format(2, m, s_per_m, s, h_per_m, h, s, h_per_m/s_per_m, h))
print(f'{3:04d}: Eine {m} hat {s_per_m} {s}n. Das sind {h_per_m:8.4f} {h}n. Eine {s} entspricht demnach {h_per_m/s_per_m:.6f} {h}n.')
```

Output:
```
0001: Eine Minute hat 60 Sekunden. Das sind   0.0167 Stunden. Eine Sekunde entspricht demnach 0.000278 Stunden.
0002: Eine Minute hat 60 Sekunden. Das sind   0.0167 Stunden. Eine Sekunde entspricht demnach 0.000278 Stunden.
0003: Eine Minute hat 60 Sekunden. Das sind   0.0167 Stunden. Eine Sekunde entspricht demnach 0.000278 Stunden.
```

<!-- grade_id: cell-10fb9b4dd9f2ec3e, points: 1, grade: true, solution: true -->
```python
# Hier eigenen Code schreiben ...

# Ausgabe für den Messbericht
for monat, werte in LAeq.items():
    max_pegel = max(werte)
    print(f"Im {monat} wurde ein maximaler Mittelungspegel von {max_pegel:.1f} dB gemessen.")
```

Output:
```
Im Januar wurde ein maximaler Mittelungspegel von 71.7 dB gemessen.
Im Februar wurde ein maximaler Mittelungspegel von 67.0 dB gemessen.
Im März wurde ein maximaler Mittelungspegel von 70.5 dB gemessen.
Im April wurde ein maximaler Mittelungspegel von 68.7 dB gemessen.
Im Mai wurde ein maximaler Mittelungspegel von 69.6 dB gemessen.
Im Juni wurde ein maximaler Mittelungspegel von 68.4 dB gemessen.
```

### 2) Richtwertüberschreitung

  * Berechnen Sie für jeden Monat die Anzahl der Tage, an denen der [zulässige Richtwert](http://www.verwaltungsvorschriften-im-internet.de/bsvwvbund_26081998_IG19980826.htm) für urbane Gebiete von 63 dB überschritten wurde. 
  * Berechnen Sie deren prozentualen Anteil an der Gesamtzahl der Werktage im jeweiligen Monat. 
  * Geben Sie das Ergebnis tabellarisch aus.
  * Vermeiden Sie auch hier die wiederholte Eingabe nahezu identischer Code-Abschnitte.
  * Speichern Sie außerdem für einen automatisierten Test die Anzahl für den Monat Juni in der Variable `anzahl_juni`. (Für die anderen Monate müssen keine eigenen Variablen erstellt werden.)

Beispiel:

    Juli    :  6 Tage (28.6 %)
    August  :  9 Tage (39.1 %)
    
<!-- Werktage Juli:21, August: 23 -->

<!-- grade_id: cell-6b290c35a4f7298f, points: 0, grade: false, solution: true -->
```python

# Hier eigenen Code schreiben ...


anzahl_juni = None # <- diese Variable sollte später die Anzahl an Juni-Tagen mit Richtwert-Überschreitung enthalten

# Richtwertüberschreitungen > 63 dB je Monat
richtwert = 63

print("\n--- Richtwertüberschreitung (> 63 dB) ---")
for monat, werte in LAeq.items():
    gesamt = len(werte)
    ueber = sum(wert > richtwert for wert in werte)
    prozent = (ueber / gesamt) * 100 if gesamt else 0.0
    print(f"{monat:8s}: {ueber:2d} Tage ({prozent:4.1f} %)")

    if monat == "Juni":
        anzahl_juni = ueber

assert anzahl_juni is not None
assert anzahl_juni < 10


```

Output:
```
--- Richtwertüberschreitung (> 63 dB) ---
Januar  :  6 Tage (27.3 %)
Februar :  6 Tage (30.0 %)
März    :  8 Tage (38.1 %)
April   :  3 Tage (15.0 %)
Mai     :  4 Tage (21.1 %)
Juni    :  9 Tage (45.0 %)
```

<!-- grade_id: cell-2d3143cf8d503b2a, points: 1, grade: true, solution: false -->
```python
# Hier ist ein Plausibilitätstest
assert anzahl_juni < 10
```

**3. Durch intensive Maßnahmen zur Lärmminderung und -vermeidung könnte der Pegel um durchschnittlich 2 dB gesenkt werden.** 

3.1) _Schätzen_ Sie basierend auf den gegebenen Werten ungefähr ab, wie viele (Werk-)Tage _im Jahr_ es noch immer zu Richtwertüberschreitungen kommen würde und geben Sie das Ergebnis aus.

<!-- grade_id: cell-a50cd0ee1f97c436, points: 0, grade: false, solution: true -->
```python
anzahl_jahr = None # <- diese Variable sollte später die geschätzte Anzahl an Tagen im Jahr mit Richtwert-Überschreitung enthalten
# Hier eigenen Code schreiben ...

reduktion = 2
grenzwert = 63 + reduktion
ueberschreitungen_6_monate = sum(1 for werte in LAeq.values() for p in werte if p > grenzwert)
anzahl_jahr = int(round((ueberschreitungen_6_monate / 6) * 12))
print(anzahl_jahr)
assert anzahl_jahr > 40

```

Output:
```
52
```

<!-- grade_id: cell-392cb039156e082d, points: 1, grade: true, solution: false -->
```python
# Hier ist ein Plausibilitätstest
assert anzahl_jahr > 40
```

3.2) Wie sieht es bei 4 dB durchschnittlicher Pegelverringerung aus?

<!-- grade_id: cell-c814317c1c47848d, points: 0, grade: false, solution: true -->
```python
anzahl_jahr = None 

# Hier eigenen Code schreiben ...

reduktion = 4
grenzwert = 63 + reduktion
ueberschreitungen_6_monate = sum(1 for werte in LAeq.values() for p in werte if p > grenzwert)
anzahl_jahr = int(round((ueberschreitungen_6_monate / 6) * 12))
print(anzahl_jahr)
assert anzahl_jahr < 40
```

Output:
```
26
```

<!-- grade_id: cell-4be286bc9f63fdaf, points: 1, grade: true, solution: false -->
```python
# Hier ist ein Plausibilitätstest
assert anzahl_jahr < 40
```

```python

```